## Problem 1 - Basic statistics of the data (*4 points*)

In this problem your task is to open and explore a NOAA weather data file using Pandas. The data file name is [data/6153237444115dat.csv](data/6153237444115dat.csv). Please perform the tasks below by writing your code into the codeblocks in each section.

**Notice**: Closely follow the instructions! For example, you should be sure to use **exactly** the same variable names mentioned in the instructions because your answers will be automatically graded, and the tests that grade your answers rely on following the same formatting or variable naming as in the instructions.

**Your score on this problem will be based on following criteria:**

- Importing the Pandas module
- Reading the data using Pandas into a variable called `data`
- Calculating a number of values from the data file used in this problem
- Including comments that explain what most lines in the code do
- Answering a couple questions at the end of the problem
- Uploading your notebook to your GitHub repository for this week's exercise  

### Part 1 (*1 point*)

Your first task is to import Pandas and read in the data file.

- Import the Pandas module (*0.5 points*)
- Read the data into a variable called `data` using Pandas (*0.5 points*)

**Important**: When reading the data, you need to tell Pandas that no-data values are specified with varying number of `*` characters. You can do this by specifying a following parameter in the Pandas `read_csv()` function: 
  
`na_values=['*', '**', '***', '****', '*****', '******']`     

<!-- grade_id: read_csv, points: 0, grade: false, solution: true -->
```python

data = None

# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

Run these tests to check that everything is set up correctly:

<!-- grade_id: problem_1_read_csv_test1, points: 0.5, grade: true, solution: false -->
```python
from nose.tools import ok_, assert_equal
import inspect
import sys

# Check that pandas module is imported
ok_('pandas' in sys.modules)
```

<!-- grade_id: problem_1_read_csv_test2, points: 0.5, grade: true, solution: false -->
```python
# Check that the variable data exists:
ok_('data' in locals())
```

### Part 2 (*1.5 points*)   

In the cell below, fill in the variables values to print answers to these questions:

- How many rows is there in the data (variable `rows`)?
- What are the column names (variable `column_names`)?
- What are the datatypes of the columns (variable `column_datatypes`)?

<!-- grade_id: cell-9cc882c07ce1e9c0, points: 0, grade: false, solution: true -->
```python
# Number of rows in the dataframe:
rows = None

# Column names:
column_names = None

# Column datatypes:
column_datatypes = None

# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_1_pandas_data_test1, points: 0.5, grade: true, solution: false -->
```python
# Print the number of rows in the dataframe:
print("There are " + str(rows) + " rows")

```

<!-- grade_id: problem_1_pandas_data_test2, points: 0.5, grade: true, solution: false -->
```python
# Print the column names:
print("The columns are: \n", column_names)


```

<!-- grade_id: problem_1_pandas_data_test3, points: 0.5, grade: true, solution: false -->
```python
# Print the column datatypes:
print("The column types are: \n", column_datatypes)


```

### Part 3 (*1.5 points*)   

In the cell below, fill in the variables values to print answers to these questions:

- What is the mean Fahrenheit temperature in the data (`temp_mean` variable calculated from the `TEMP` column)?
- What is the standard deviation of the Maximum temperature (`temp_max_std` variable calculated from the `MAX` column)?
- How many unique stations exists in the data (`station_count` variable calcualted from the `USAF` column)?

<!-- grade_id: basic_statistics, points: 0, grade: false, solution: true -->
```python
# Calculate mean temperature
temp_mean = None

# Calculate standard deviaton of the maximum temperature:
temp_max_std = None

#Calculate number of unique stations:
station_count = None

# REPLACE THE ERROR BELOW WITH YOUR OWN CODE
raise NotImplementedError()
```

<!-- grade_id: problem_1_basic_statistics_test1, points: 0.5, grade: true, solution: false -->
```python
# Check mean temperature value
print("The mean temperature in Fahrenheit is ", round(temp_mean,1))

```

<!-- grade_id: problem_1_basic_statistics_test2, points: 0.5, grade: true, solution: false -->
```python
# Check standard deviation value
print("The standard deviation of maximum temperature is ", round(temp_max_std, 1))

```

<!-- grade_id: problem_1_basic_statistics_test, points: 0.5, grade: true, solution: false -->
```python
# Check number of stations value
print("The number of unique stations is ", station_count)

```

### Problem 1 summary

In the [Exercise 5 summary notebook](Exercise-5-summary.ipynb) you can find a few additional points to consider and one final question for Problem 1. Please answer that question in [that notebook](Exercise-5-summary.ipynb).

### On to Problem 2

Now you can continue to [Problem 2: Data manipulation and selection](Exercise-5-problem-2.ipynb)